# Predicting beer type

This work is to predict what type of beer each is based on the characteristics of that beer.

The data I will use here come from a publicly-available [Kaggle dataset on craft beer](https://www.kaggle.com/nickhould/craft-cans).

# Part I : Data, Wrangling, & EDA

<!-- grade_id: cell-1fa5c018cbfab919, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import numpy as np
from sklearn.svm import SVC
from sklearn.metrics import confusion_matrix, classification_report, precision_recall_fscore_support
```

<!-- grade_id: cell-d8b87114e650e607, points: 0.125, grade: true, solution: false -->
```python
assert pd
assert np
assert SVC
assert confusion_matrix
assert classification_report
assert precision_recall_fscore_support
```

<!-- grade_id: cell-d10db4e6f0c02976, points: 0, grade: false, solution: true -->
```python
breweries = pd.read_csv("data/breweries.csv")
beers = pd.read_csv("data/beers.csv")
```

<!-- grade_id: cell-a3a4511bf1f57302, points: 0.125, grade: true, solution: false -->
```python
assert breweries.shape == (558, 4)
assert beers.shape == (2410, 8)
```

```python
breweries.head()
```

Output:
```
Unnamed: 0                       name           city state
0           0         NorthGate Brewing     Minneapolis    MN
1           1  Against the Grain Brewery     Louisville    KY
2           2   Jack's Abby Craft Lagers     Framingham    MA
3           3  Mike Hess Brewing Company      San Diego    CA
4           4    Fort Point Beer Company  San Francisco    CA
```

```python
beers.head()
```

Output:
```
Unnamed: 0    abv  ibu    id                 name  \
0           0  0.050  NaN  1436             Pub Beer   
1           1  0.066  NaN  2265          Devil's Cup   
2           2  0.071  NaN  2264  Rise of the Phoenix   
3           3  0.090  NaN  2263             Sinister   
4           4  0.075  NaN  2262        Sex and Candy   

                            style  brewery_id  ounces  
0             American Pale Lager         408    12.0  
1         American Pale Ale (APA)         177    12.0  
2                    American IPA         177    12.0  
3  American Double / Imperial IPA         177    12.0  
4                    American IPA         177    12.0
```

<!-- grade_id: cell-d97505cb97ec42e8, points: 0, grade: false, solution: true -->
```python
null_beers = beers.isnull().sum()
```

<!-- grade_id: cell-388ac707168e1d78, points: 0, grade: false, solution: true -->
```python
beers = beers.dropna(subset=['style', 'abv', 'ibu'])
```

<!-- grade_id: cell-3f94ee193183fa7d, points: 0, grade: false, solution: true -->
```python
beer_df = beers.merge(breweries, how = "left")
beer_df.head()
```

Output:
```
Unnamed: 0    abv   ibu    id                                  name  \
0          14  0.061  60.0  1979                          Bitter Bitch   
1          21  0.099  92.0  1036                         Lower De Boom   
2          22  0.079  45.0  1024                         Fireside Chat   
3          24  0.044  42.0   876                       Bitter American   
4          25  0.049  17.0   802  Hell or High Watermelon Wheat (2009)   

                     style  brewery_id  ounces city state  
0  American Pale Ale (APA)         177    12.0  NaN   NaN  
1      American Barleywine         368     8.4  NaN   NaN  
2            Winter Warmer         368    12.0  NaN   NaN  
3  American Pale Ale (APA)         368    12.0  NaN   NaN  
4   Fruit / Vegetable Beer         368    12.0  NaN   NaN
```

```python
beer_df.describe()
```

Output:
```
Unnamed: 0          abv          ibu           id   brewery_id  \
count  1403.000000  1403.000000  1403.000000  1403.000000  1403.000000   
mean   1241.128297     0.059919    42.739843  1413.888810   223.375624   
std     691.675612     0.013585    25.962692   757.572191   150.387510   
min      14.000000     0.027000     4.000000     1.000000     0.000000   
25%     681.500000     0.050000    21.000000   771.000000    95.500000   
50%    1228.000000     0.057000    35.000000  1435.000000   198.000000   
75%    1864.500000     0.068000    64.000000  2068.500000   350.000000   
max    2408.000000     0.125000   138.000000  2692.000000   546.000000   

            ounces  
count  1403.000000  
mean     13.510264  
std       2.254112  
min       8.400000  
25%      12.000000  
50%      12.000000  
75%      16.000000  
max      32.000000
```

<!-- grade_id: cell-f6596e92382da73f, points: 0, grade: false, solution: true -->
```python
beer_counts = beer_df["style"].value_counts()
print(beer_counts)
```

Output:
```
American IPA                          301
American Pale Ale (APA)               153
American Amber / Red Ale               77
American Double / Imperial IPA         75
American Blonde Ale                    61
                                     ... 
Roggenbier                              1
Smoked Beer                             1
Euro Pale Lager                         1
Other                                   1
American Double / Imperial Pilsner      1
Name: style, Length: 90, dtype: int64
```

<!-- grade_id: cell-48c99ba398ca62d3, points: 0, grade: false, solution: true -->
```python
styles = list(beer_counts.index[0: 4])
beer_df = beer_df[beer_df["style"].isin(styles)]
```

# Part II : Prediction Model

<!-- grade_id: cell-f165209e70ea6449, points: 0, grade: false, solution: true -->
```python
import random
num_training = int(0.8 * beer_df.shape[0])
num_testing = beer_df.shape[0] - num_training
```

<!-- grade_id: cell-09abd08ce060707e, points: 0, grade: false, solution: true -->
```python
beer_X = beer_df[["abv", "ibu"]]
beer_Y = np.array(beer_df["style"])
```

<!-- grade_id: cell-fd8962674731f8e3, points: 0, grade: false, solution: true -->
```python
beer_train_X = beer_X.iloc[: num_training]
beer_test_X = beer_X.iloc[num_training: ]
beer_train_Y = beer_Y[: num_training]
beer_test_Y = beer_Y[num_training: ]
```

```python
def train_SVM(X, y, kernel='linear'):
    clf = SVC(kernel=kernel)
    clf.fit(X, y)
    
    return clf
```

<!-- grade_id: cell-c206bff1f25a305d, points: 0, grade: false, solution: true -->
```python
beer_clf = train_SVM(beer_train_X, beer_train_Y)
```

<!-- grade_id: cell-2a2fef981469581b, points: 0, grade: false, solution: true -->
```python
beer_predicted_train_Y = beer_clf.predict(beer_train_X)
beer_predicted_test_Y = beer_clf.predict(beer_test_X)
```

# Part III : Model Assessment

<!-- grade_id: cell-5cc9c9d3cca4e6ce, points: 0, grade: false, solution: true -->
```python
class_report_train = classification_report(beer_train_Y, beer_predicted_train_Y)
print(class_report_train)
```

Output:
```
precision    recall  f1-score   support

      American Amber / Red Ale       0.82      0.45      0.58        69
American Double / Imperial IPA       0.76      0.25      0.37        53
                  American IPA       0.69      0.84      0.76       236
       American Pale Ale (APA)       0.57      0.64      0.60       126

                      accuracy                           0.67       484
                     macro avg       0.71      0.54      0.58       484
                  weighted avg       0.69      0.67      0.65       484
```

<!-- grade_id: cell-82e29dc9780e9d89, points: 0, grade: false, solution: true -->
```python
class_report_test = classification_report(beer_test_Y, beer_predicted_test_Y)
print(class_report_test)
```

Output:
```
precision    recall  f1-score   support

      American Amber / Red Ale       0.62      0.62      0.62         8
American Double / Imperial IPA       0.78      0.32      0.45        22
                  American IPA       0.70      0.72      0.71        65
       American Pale Ale (APA)       0.55      0.78      0.65        27

                      accuracy                           0.66       122
                     macro avg       0.66      0.61      0.61       122
                  weighted avg       0.68      0.66      0.64       122
```

<!-- grade_id: cell-f0a1f8d304bd71ae, points: 0, grade: false, solution: true -->
```python
conf_mat_train = confusion_matrix(beer_train_Y, beer_predicted_train_Y)
print(conf_mat_train)
```

Output:
```
[[ 31   1  10  27]
 [  0  13  40   0]
 [  0   3 198  35]
 [  7   0  38  81]]
```

<!-- grade_id: cell-dcfae2fde6f81ba9, points: 0, grade: false, solution: true -->
```python
conf_mat_test = confusion_matrix(beer_test_Y, beer_predicted_test_Y)
print(conf_mat_test)
```

Output:
```
[[ 5  0  2  1]
 [ 1  7 14  0]
 [ 0  2 47 16]
 [ 2  0  4 21]]
```